Before you turn this problem in, make sure everything runs as expected. **Restart the kernel and run all cells** (in the menubar, select Kernel$\rightarrow$Restart & Run All).

Make sure you fill in any place that says `YOUR CODE HERE ???` or "YOUR ANSWER HERE". Code cells are autograded, and if the test cells don't run all the way to the line that  prints "Tests Passed (# points)", you will not recieve any point for that test.

# CSCI 3155: Assignment 9

Topics: 
- Type Inference
- Basics of Objects
- Variance Annotations.

Readings: Notes posted on moodle.  Chapters from Odersky

- Chapter 4 (revise)
- Chapter 10
- Chapter 11 (take note of Any, AnyVal, Null, Nothing)
- Chapter 12 
- Chapter 19

__Name__: Theodore Margoles

```scala
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (20 points): Objects in Scala

### A (12 points)
Consider the class below for an employee record in an organization that has three fields `name`, `eid` (employee id) and `salary`.

1) As given, scala will complain that class employee cannot extend the abstract class People. Add the missing methods to enable the class Employee to properly extend People

2) Write a factory pattern that takes in a formatted string of the form "John Doe,29585,33110" consisting of a string and two numbers, to extract the name and salaries for each employee. You may want to use `scala.util.matching.Regex` [https://www.scala-lang.org/api/2.9.2/scala/util/matching/Regex.html]
If the input fails to match the pattern, an `IllegalArgumentException` must be thrown

Companion objects are covered in scala book Chapter 4.3. Also lookup here: https://alvinalexander.com/scala/factory-pattern-in-scala-design-patterns

<!-- grade_id: cell-f80a491b407a7ab7, points: 0, grade: false, solution: true -->
```scala
import scala.util.matching.Regex

abstract class People {
    /* I need the ability to convert every person to a string */
    def toString: String
    /* I need the ability to check if two people are the same */
    def equals (p: People): Boolean = {
        if(this.toString == p.toString) true else false
    }
}

/* 
  Note that for an employee to be equal to an object of type People, that object
  must be an employee with matching eid number 
*/

class Employee( val name: String, val eid: Int, var salary: Int) extends People {
    override def toString: String =  name + " with salary = " + salary.toString
    
    
}


object Employee {
    def apply(formattedInput: String): Employee = {
        val nums:List[Char] = List('1', '2', '3', '4', '5', '6', '7', '8', '9', '0')
        //def Num(c:Char, l:List[Char]):Boolean ={if(l.contains(c)) true else false}

        /* 
        Input format must be 
        (full name with possible spaces)[optional whitespaces],[optional whitespaces](employee ID number)[optional whitespaces],[optional whitespaces](employee salary)
        */
        val name = (formattedInput.filterNot(x => nums.contains(x))).trim
        val numchars = (formattedInput.filter(x => nums.contains(x) || x ==  ' ')).trim
        val sarr:Array[String] = numchars.split(',')
        new Employee(name, sarr(0).toInt, sarr(1).toInt)
    }
}
```

Output:
```
[32mimport [39m[36mscala.util.matching.Regex

[39m
defined [32mclass[39m [36mPeople[39m
defined [32mclass[39m [36mEmployee[39m
defined [32mobject[39m [36mEmployee[39m
```

<!-- grade_id: cell-e1ec3bf75d6f6de5, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
println(e1.toString)
assert(e1.name == "Jonathan Doe", "Failed")
assert(e1.eid == 12893, "Failed")
assert(e1.salary == 39813, "Failed")
assert(e1.equals(new Employee("whoknows", 12893, 39813)))
passed(4)
//END TEST
```

Output:
```
Jonathan Doe, , with salary = 39813
  $sess.cmd6Wrapper.<init>(cmd6.sc:128)
  $sess.cmd6$.<init>(cmd6.sc:76)
  $sess.cmd6$.<clinit>(cmd6.sc:-1)
```

<!-- grade_id: cell-9d823f9842fd95a7, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
try {
    val e2 = Employee("Badly Formatted String, 29812, $3102")
    assert(false, "Failed, the string is badly formatted but your code is OK with it.")
} catch {
    case e => println(s"Expected behavior seen")
}
passed(4)
//END TEST
```

Output:
```
Expected behavior seen

*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-00d1907d97c7da69, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Janet Maria Catherine Doe    , 12894   , 121331")
println(e1.toString)
assert(e1.name == "Janet Maria Catherine Doe", "Failed")
assert(e1.eid == 12894, "Failed")
assert(e1.salary == 121331, "Failed")
assert(e1.equals(new Employee("random", 12894, 121332)))
passed(4)
//END TEST
```

Output:
```
$sess.cmd4Wrapper.<init>(cmd4.sc:128)
  $sess.cmd4$.<init>(cmd4.sc:76)
  $sess.cmd4$.<clinit>(cmd4.sc:-1)
```

### B (8 points)
We wish to add a new class of people to our system. These are "Contractor" with the information associated should involve their name (string), contract ID (Integer) and contract supervisor (Employee).

Write a new class named `Contractor` whose constructor should take in the name of the person (String), their ID number (Integer) and Supervisor (Employee) and extend from People. 

A contractor object is equal to an object p of type People, only if p is actually an instance of contractor, and if their names are the same and they have the same contract ID number.

<!-- grade_id: cell-97ab958367cfbd0a, points: 0, grade: false, solution: true -->
```scala
class Contractor(name:String, contractID:Int, contractSupervisor:Employee) extends People{
    override def toString =  name + contractID + contractSupervisor.toString
}
```

Output:
```
defined [32mclass[39m [36mContractor[39m
```

<!-- grade_id: cell-82cc00e175a12efe, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
val c1 = new Contractor("John Contract Worker", 12093, e1)
assert( !c1.equals(e1), "Failed")
assert( c1.equals(c1), "Failed")
passed(4)
//END TEST
```

Output:
```
$sess.cmd11Wrapper.<init>(cmd11.sc:128)
  $sess.cmd11$.<init>(cmd11.sc:80)
  $sess.cmd11$.<clinit>(cmd11.sc:-1)
```

<!-- grade_id: cell-62fa2e0b66f58862, points: 4, grade: true, solution: false -->
```scala
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
val e2 = Employee("Jane Doe, 12894, 43814")
val c1 = new Contractor("John Contract Worker", 12093, e2)
val c2 = new Contractor("John Contract Worker", 12093, e1)
assert( c1.equals(c2), "Failed")
passed(4)
//END TEST
```

Output:
```
$sess.cmd12Wrapper.<init>(cmd12.sc:132)
  $sess.cmd12$.<init>(cmd12.sc:80)
  $sess.cmd12$.<clinit>(cmd12.sc:-1)
```

## Problem 2 (40 points): Traits in Scala

For this problem, we have defined two traits: `NumberOfLegs` that helps us define how many legs a given animal has and `WarmBlooded` that applies to warm blooded animals.

```scala
abstract class Animal 

trait NumberOfLegs {
    val nLegs: Int
    def getNumberOfLegs: Int = nLegs
}

trait WarmBlooded extends Animal {
    val bodyTempMaintained: Double
    def getBodyTemp: Double = bodyTempMaintained 
}
```

Output:
```
defined [32mclass[39m [36mAnimal[39m
defined [32mtrait[39m [36mNumberOfLegs[39m
defined [32mtrait[39m [36mWarmBlooded[39m
```

### A (5 points)

Define a class `Human` that inherits from `Animal` and mixes in the traits `NumberOfLegs` with `nLegs = 2` and `WarmBlooded` with `bodyTempMaintained = 98`. The class `Human` must take in a parameter called `name` of type `String` and implement a `getName` method without any parameters.

<!-- grade_id: cell-e7a4a9275bd24e24, points: 0, grade: false, solution: true -->
```scala
class Human(name:String) extends Animal with NumberOfLegs with WarmBlooded{
   override val nLegs = 2
   override val bodyTempMaintained = 98.0
   def getName = name
}
```

Output:
```
defined [32mclass[39m [36mHuman[39m
```

<!-- grade_id: cell-f9b7ae8902326372, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val t1 = new Human("Jane Smith")
assert(t1.getNumberOfLegs == 2, "Your human does not have two legs")
assert(t1.bodyTempMaintained == 98.0, "Your human does not maintain a body temp of 98")
assert(t1.getBodyTemp == 98.0, "Your human's getBodyTemp Function is not working")
assert(t1.getName == "Jane Smith", "Your human's name is not setting correctly")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36mt1[39m: [32mHuman[39m = $sess.cmd15Wrapper$Helper$Human@990efe1
```

### B (5 points)
Define a class named `Table` that mixes in the trait `NumberOfLegs` with `nLegs = 4`. 

_Note:_ `Table` cannot mixin the trait `WarmBlooded` (make sure you understand why that is as a quick check of concepts).

<!-- grade_id: cell-8145df690b6f6588, points: 0, grade: false, solution: true -->
```scala
class Table extends NumberOfLegs{
    override val nLegs = 4
}
```

Output:
```
defined [32mclass[39m [36mTable[39m
```

<!-- grade_id: cell-d34b29b062b93894, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
val tbl = new Table()
assert(tbl.getNumberOfLegs == 4, "A Table must have four legs")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
[36mtbl[39m: [32mTable[39m = $sess.cmd17Wrapper$Helper$Table@54ae2f71
```

### C (15 points)
Define a class called `Menagerie` with type parameter `T` such that 

- T must always be a derived type of Animal with the trait NumberOfLegs (In scala you can say `T <: A with B` to indicate that a type parameter T must derive from class A with trait B applied).

- `Menagerie` has a list of objects of type `T`.

- The user must be able to instantiate an object of type Menagerie without providing any class parameters: `new Menagerie[T]`

- The class Menagerie has the method `addAnimalToMenagerie` that adds an object of type `T` to menagerie and returns a new Menagerie.

- Implement a method `addMultipleAnimalsToMenagerie` that adds more than one animal. Lookup syntax below.
 
 
Lookup the syntax for varargs or repeated parameters (T*) here: https://alvinalexander.com/scala/how-to-define-methods-variable-arguments-varargs-fields. Also in Chapter 8.8 of Odersky book.


- Add a method `totalLegs: Int` that sums up how many legs the animals in the menagerie have in total.

~~~
class Menagerie[T...] () {
   // ... Extra Stuff that you may want to add ...
   def addAnimalToMenagerie (a: T) : Menagerie[T]
   def addMultipleAnimalsToMenagerie(animals: T*): Menagerie[T]
   def totalLegs:Int
}
~~~


<!-- grade_id: cell-44221c4360112bc9, points: 0, grade: false, solution: true -->
```scala
class Menagerie[T <: Animal with NumberOfLegs](objects:List[T]){
    def this() = this(List[T]())
    
    def addAnimalToMenagerie(a : T) : Menagerie[T] = {
        new Menagerie[T](a::objects)
    }
    
    def addMultipleAnimalsToMenagerie(animals: T*): Menagerie[T] = {
        new Menagerie(animals.foldLeft(List[T]())((acc, el) => el::acc))
    }
    
    def totalLegs : Int = {
        objects.foldLeft(0)((acc, el) => acc + el.getNumberOfLegs)
    }
    
}
```

Output:
```
defined [32mclass[39m [36mMenagerie[39m
```

```scala
//BEGIN TEST SETUP
class Donkey extends Animal with NumberOfLegs with WarmBlooded {
    val nLegs:Int = 4
    val bodyTempMaintained = 96.0
}

class Lion extends Animal with NumberOfLegs {
    val nLegs:Int = 4
}

class Mule extends Donkey

val d1 = new Donkey()
val d2 = new Donkey()
val d3 = new Donkey()
val d4 = new Mule()
//END TEST SETUP
```

Output:
```
defined [32mclass[39m [36mDonkey[39m
defined [32mclass[39m [36mLion[39m
defined [32mclass[39m [36mMule[39m
[36md1[39m: [32mDonkey[39m = $sess.cmd20Wrapper$Helper$Donkey@544d80da
[36md2[39m: [32mDonkey[39m = $sess.cmd20Wrapper$Helper$Donkey@3440277e
[36md3[39m: [32mDonkey[39m = $sess.cmd20Wrapper$Helper$Donkey@5117bcaa
[36md4[39m: [32mMule[39m = $sess.cmd20Wrapper$Helper$Mule@6d79844e
```

<!-- grade_id: cell-d70b45ba2426c6a6, points: 3, grade: true, solution: false -->
```scala
//BEGIN TEST
val m1 = new Menagerie[Donkey]
passed(3)
//END TEST
```

Output:
```
*** Tests Passed (3 points) ***
[36mm1[39m: [32mMenagerie[39m[[32mDonkey[39m] = $sess.cmd19Wrapper$Helper$Menagerie@f1db976
```

<!-- grade_id: cell-0f3cd5e1e161c390, points: 6, grade: true, solution: false -->
```scala
//BEGIN TEST
val m1 = new Menagerie[Donkey]
val m2 = m1.addAnimalToMenagerie(d1)
val n = m2.totalLegs
assert(n == 4, s"1 quadraped = 4 legs. You have $n legs : are you missing legs? Too many legs?")
passed(6)
//END TEST
```

Output:
```
*** Tests Passed (6 points) ***
[36mm1[39m: [32mMenagerie[39m[[32mDonkey[39m] = $sess.cmd19Wrapper$Helper$Menagerie@6ccb6d1e
[36mm2[39m: [32mMenagerie[39m[[32mDonkey[39m] = $sess.cmd19Wrapper$Helper$Menagerie@150bb218
[36mn[39m: [32mInt[39m = [32m4[39m
```

<!-- grade_id: cell-9ca767949c491818, points: 6, grade: true, solution: false -->
```scala
//BEGIN TEST
val m1 = new Menagerie[Donkey]
val m2 = m1.addMultipleAnimalsToMenagerie(d1, d2, d3, d4)
val n = m2.totalLegs
assert(n == 16, s"4 quadrapeds = 16 legs. You have $n legs : are you missing legs? Too many legs?")
passed(6)
//END TEST
```

Output:
```
*** Tests Passed (6 points) ***
[36mm1[39m: [32mMenagerie[39m[[32mDonkey[39m] = $sess.cmd19Wrapper$Helper$Menagerie@32a390b1
[36mm2[39m: [32mMenagerie[39m[[32mDonkey[39m] = $sess.cmd19Wrapper$Helper$Menagerie@4791b8d4
[36mn[39m: [32mInt[39m = [32m16[39m
```

### D (15 points)
Currently a `Menagerie` as defined above is Invariant. Make it covariant so that  subtyping relations such as `Menagerie[Mule] <: Menagerie[Donkey] <: Menagerie[Animal]`.

<!-- grade_id: cell-3dcb8201b939b421, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
???
```

<!-- grade_id: cell-96824e578d275d9c, points: 15, grade: true, solution: false -->
```scala
//BEGIN TEST 
class Donkey extends Animal with NumberOfLegs with WarmBlooded {
    val nLegs:Int = 4
    val bodyTempMaintained = 96.0
}

class Mule extends Donkey

val d1 = new Donkey()
val d2 = new Donkey()
val d3 = new Donkey()
val mu1 = new Mule()
val m1 = new Menagerie[Donkey]
// Inheritance among Menageries:
val m2: Menagerie[Donkey] = new Menagerie[Mule]
// Strongest type selection:
val m3: Menagerie[Donkey] = m1.addMultipleAnimalsToMenagerie(d1, d2, d3, mu1)

print("Successfully added Lions and Donkeys to the same menagerie!")

passed(15)
//END TEST
```